---
title       : TaschenrechneR
description : R als Taschenrechner nutzen
attachments :


--- type:NormalExercise lang:r xp:100 skills:1
## Einfache Arithmetik

In the previous exercise, you saw a dataset about movies. In this exercise, we'll have a look at yet another dataset about movies!

A dataset with a selection of movies, `movie_selection`, is available in the workspace.

*** =instructions
-  Beginnen wir mit einfacher Arithmetik. Bestimmen Sie die Summe von 3 und 5.

*** =hint
-  Tippen Sie wie im Taschenrechner 3+5 ein.

*** =pre_exercise_code
```{r}
# Pre-load a package in the workspace

# You can prepare the data before the student starts:

# You can also clean up data so that it's not available in the student's workspace anymore:
```

*** =sample_code
```{r}

```

*** =solution
```{r}
3 + 5
```

*** =sct
```{r}
# The sct section defines the Submission Correctness Tests (SCTs) used to
# evaluate the student's response. All functions used here are defined in the 
# testwhat R package. Documentation can also be found at github.com/datacamp/testwhat/wiki

# Test whether the function str is called with the correct argument, object
# If it is not called, print something informative
# If it is called, but called incorrectly, print something else
#test_function("str", args = "object",
#              not_called_msg = "You didn't call `str()`!",
#              incorrect_msg = "You didn't call `str(object = ...)` with the correct argument, `object`.")

# Test the object, good_movies
# Notice that we didn't define any feedback here, this will cause automatically 
# generated feedback to be given to the student in case of an incorrect submission
#test_object("good_movies")

# Test whether the student correctly used plot()
# Again, we use the automatically generated feedback here
#test_function("plot", args = "x")
#test_function("plot", args = "y")
#test_function("plot", args = "col")

# Alternativeley, you can use test_function() like this
# test_function("plot", args = c("x", "y", "col"))

# It's always smart to include the following line of code at the end of your SCTs
# It will check whether executing the student's code resulted in an error, 
# and if so, will cause the exercise to fail
test_error()

# Final message the student will see upon completing the exercise
success_msg("Good work!")
```
